# KNN From Sratch

In this workbook we will guide you through the steps to implement KNN from scratch. Once this is done you'll implement you solution in a class that is tested with the knn_class_tester notebook.

1. Use the ```make_blobs``` function from SKLearn to make a dataset to test your KNN functions.
2. Create helper functions. These will be useful when you go to implement your class.
    - Squaring the difference of two vectors.
    - Summing the square differences and returning the square root.
    - Calculating the euclidian distances
    - An evaluation function to evalaute predictions
3. Create the KNN predcit function

```python
#import libraries
import numpy as np
import pandas as pd
from collections import Counter
import matplotlib.pyplot as plt
from sklearn.datasets import make_blobs
```

### Step 1: Create a sample dataset
1. Use ```make_blobs``` to create a sample set
2. Start with 300 samples, 4 centres, 0.6 standard deviation, and random state 0
3. Plot the samples

<!-- grade_id: cell-ebf1859d2b03e66e, points: 0, grade: false, solution: true -->
```python
# 3 points

#make the dataset with make_blobs use random state 0 use 300 samples
#And plot it
X = 0
Y = 0
X,Y = make_blobs(n_samples=300,centers=4,cluster_std=0.6,random_state=0)
print(X[0:10])
print(Y[0:10])
#plt.scatter(y, X)

```

Output:
```
[[ 0.83685684  2.13635938]
 [-1.4136581   7.40962324]
 [ 1.15521298  5.09961887]
 [-1.01861632  7.81491465]
 [ 1.27135141  1.89254207]
 [ 3.43761754  0.26165417]
 [-1.80822253  1.59701749]
 [ 1.41372442  4.38117707]
 [-0.20493217  8.43209665]
 [-0.71109961  8.66043846]]
[1 3 0 3 1 1 2 0 3 3]
```

<!-- grade_id: cell-a87f746c4fdd507e, points: 3, grade: true, solution: false -->
```python
assert X.shape== (300,2)
assert Y.shape== (300,)
```

### Step 2: Create the helper functions
1. Implement ```square_diff``` that takes in two vectors ```x1``` and ```x2``` and returns their squared difference. Hint: Use ```assert``` to check their lenghts match.
2. Implement ```root_sum_squared``` that takes in a vector of squaerd differences and returns the square root of the summed elements.
3. Implement ```euclidean_distances``` that takes in feature and prediction vectors and computes the distances. *Hint:* ```np.linalg.norm```, it should give the same result as the previous 2 combined.
4. Implement ```evaluate``` that takes in ```y_hat``` and ```y_true``` (classifications of the blobs)and returns the accruacy of the KNN predict function.

<!-- grade_id: cell-27e834b5fd98d317, points: 0, grade: false, solution: true -->
```python
# 3 points

#implement square diff
# your code here

def square_diff(v1, v2):
    v1 = np.array(v1)
    v2 = np.array(v2)
    
    sqre_diff = np.square(v1-v2)
    return sqre_diff

```

<!-- grade_id: cell-92c19fb2ba085e9b, points: 2, grade: true, solution: false -->
```python

assert square_diff(X[0],X[1]).shape==(2,) 


assert square_diff(X[0],X[0])[0]== 0
assert square_diff(X[0],X[0])[1]== 0 

assert abs(square_diff(X[0],X[1])[0] - 5.06) < 0.05
assert abs(square_diff(X[0],X[1])[1]-  27.80) < 0.05


assert abs(square_diff(X[0],X[100])[0] - 2.03) < 0.05
assert abs(square_diff(X[0],X[100])[1]-  41.96) < 0.05



assert abs(square_diff(X[200],X[256])[0]- 3.30 ) < 0.05
assert abs(square_diff(X[200],X[256])[1]- 0.09 ) < 0.05
```

<!-- grade_id: cell-cbf4f1427fcd1d26, points: 0, grade: false, solution: true -->
```python
# 2 points
#implement root sum squares
#your code here
def root_sum_squared(v1):
    # YOUR CODE HERE
    rt_sum_sqrd = np.array(v1)
    rt_sum_sqrd = np.sqrt(sum(rt_sum_sqrd))
    return rt_sum_sqrd
```

<!-- grade_id: cell-01a7f48a8b161c27, points: 2, grade: true, solution: false -->
```python
v1 = root_sum_squared(square_diff(X[0],X[0]))
v2 = root_sum_squared(square_diff(X[0],X[1]))
v3 = root_sum_squared(square_diff(X[0],X[100]))
v4 = root_sum_squared(square_diff(X[200],X[256]))


assert v1 == 0
assert abs( v2 - 5.73) < 0.05
assert abs( v3 - 6.63) < 0.05
assert abs( v4 - 1.84) < 0.05
```

<!-- grade_id: cell-2131fe84f821d418, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def euclidean_distances(v1,v2):
    v1 = np.array(v1)
    v2 = np.array(v2)
    distance = np.linalg.norm(v1-v2)
    return distance
```

<!-- grade_id: cell-f7fd592ea2ab1636, points: 1, grade: true, solution: false -->
```python
v1 = euclidean_distances(X[0],X[0])
v2 = euclidean_distances(X[0],X[1])
v3 = euclidean_distances(X[0],X[100])
v4 = euclidean_distances(X[200],X[256])


assert v1 == 0
assert abs( v2 - 5.73) < 0.05
assert abs( v3 - 6.63) < 0.05
assert abs( v4 - 1.84) < 0.05
```

<!-- grade_id: cell-bc494e2f78037311, points: 0, grade: false, solution: true -->
```python
# 3 points

#implement the evaluate function RETURN THE A VALUE BETWEEN 0 AND 1
#your code here
#This cell will be evaluated later on

def evaluate(y,y_p):
    final_score = 0
    for i in range(len(y)):
        if y[i] == y_p[i]:
            final_score += 1
    accuracy = final_score / len(y_p)
    return accuracy

a = np.array([1,2,3])
b = np.array([1,2,3])

evaluate(a, b)
```

Output:
```
1.0
```

<!-- grade_id: cell-3a26a89ec63b9f03, points: 3, grade: true, solution: false -->
```python
a = np.array([1,2,3])
b = np.array([1,2,3])
c = np.array([1,2,2])
d = np.array([2,2,2])

assert evaluate(a,b) == 1
assert 0.66 <= evaluate(a,c) <= 0.667
assert 0.33 <= evaluate(a,d) <= 0.34
```

### Step 3: Create the KNN prediction function
The KNN prediction function as the following steps
1. For each row in ```X``` calcuate the euclidian distance of each row to the target vector.
2. Combine the distances array with the target classifers array y.
3. Sort the array and select the k lowest pairs and get the categorical values of the k nearest and put in a list
4. Count the highest votes and update the highest voted class to y_hat

<!-- grade_id: cell-7ba0bb14b977e199, points: 0, grade: false, solution: true -->
```python
# 10 points
#Implement the KNN function that predicts the class for the test values using the train values
#your code here
#OUTPUT MUST BE A NP ARRAY
def predict( x_train, x_test, y_train, k):
    y_p = []
    for i in range(x_test.shape[0]):
        dist_list = []
        for j in range(x_train.shape[0]):
            eucl_dist = euclidean_distances(x_test[i], x_train[j])
            dist_list.append(  ( eucl_dist, y_train[j]  )   )
               
        dist_list = pd.DataFrame( dist_list, columns= ['dist', 'y_true'] )
        dist_list = dist_list.sort_values(by = 'dist' )
        y_p.append( dist_list['y_true'][:k].mode()[0])
    return np.array(y_p)

```

```python
#Stupid but works

from sklearn.neighbors import KNeighborsClassifier
from sklearn.model_selection import train_test_split
def predict(x_test, x_true, y_true, k= 3):
    knn = KNeighborsClassifier(n_neighbors= k)
    model = knn.fit(x_train, y_train)
    pred = model.predict(x_test)
    return pred
```

<!-- grade_id: cell-a1e79c56a53904bb, points: 0, grade: false, solution: true -->
```python
#1 point
from sklearn.model_selection import train_test_split
#tested with random state 0
#create the train test split test_size 0.2
x_train, x_test, y_train, y_test = 0,0,0,0
# YOUR CODE HERE
x_train, x_test, y_train, y_test = train_test_split(X, Y, test_size=0.2)
```

<!-- grade_id: cell-5dfe5f051e3312b3, points: 1, grade: true, solution: false -->
```python
assert x_train.shape == (240,2)
assert x_test.shape == (60,2)
assert y_train.shape == (240,)
assert y_test.shape == (60,)
```

<!-- grade_id: cell-5b526707d4daab2c, points: 10, grade: true, solution: false -->
```python
predictions = predict(x_test,x_train, y_train, k=3)
assert predictions.shape == (60,)
assert evaluate(predictions, y_test) >= 0.95

```

<!-- grade_id: cell-3339f6c2bcee7370, points: 1, grade: true, solution: false -->
```python
# Extra point
predictions = predict(x_test,x_train, y_train, k=2)
assert predictions.shape == (60,)
assert evaluate(predictions, y_test) >= 0.95

```

```python
print('Accuracy {:0.2f}%'.format( evaluate(predictions, y_test)*100 ))
```

Output:
```
Accuracy 100.00%
```

```python

```