***Support Vector Machines***
    Using SVM with Python to predict whether a breast cancer tumor is malignant or benign. Exercise uses numpy, pandas, and scikitlearn and utilizes train test split, SVM, SVC,  and GridSearch to identify the best parameters for prediction.

```python
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import seaborn as sns
%matplotlib inline
```

```python

```

<!-- grade_id: cell-bbc3dc112799b91f, points: 0, grade: false, solution: true -->
```python
# 1 point

# import and load the breast cancer dataset 
from sklearn.datasets import load_breast_cancer
# output a variable where the data is loaded straight from sklearn
cancer =load_breast_cancer()
print(dir(cancer))





```

Output:
```
['DESCR', 'data', 'feature_names', 'filename', 'frame', 'target', 'target_names']
```

<!-- grade_id: cell-89604aed63202fe9, points: 1, grade: true, solution: false -->
```python
assert cancer.data.shape == (569,30)
```

```python
print(cancer['DESCR'])
```

Output:
```
.. _breast_cancer_dataset:

Breast cancer wisconsin (diagnostic) dataset
--------------------------------------------

**Data Set Characteristics:**

    :Number of Instances: 569

    :Number of Attributes: 30 numeric, predictive attributes and the class

    :Attribute Information:
        - radius (mean of distances from center to points on the perimeter)
        - texture (standard deviation of gray-scale values)
        - perimeter
        - area
        - smoothness (local variation in radius lengths)
        - compactness (perimeter^2 / area - 1.0)
        - concavity (severity of concave portions of the contour)
        - concave points (number of concave portions of the contour)
        - symmetry
        - fractal dimension ("coastline approximation" - 1)

        The mean, standard error, and "worst" or largest (mean of the three
        worst/largest values) of these features were computed for each image,
        resulting in 30 features.  For instance, field 0 is Mean Radi
```

```python
cancer['feature_names']
```

Output:
```
array(['mean radius', 'mean texture', 'mean perimeter', 'mean area',
       'mean smoothness', 'mean compactness', 'mean concavity',
       'mean concave points', 'mean symmetry', 'mean fractal dimension',
       'radius error', 'texture error', 'perimeter error', 'area error',
       'smoothness error', 'compactness error', 'concavity error',
       'concave points error', 'symmetry error',
       'fractal dimension error', 'worst radius', 'worst texture',
       'worst perimeter', 'worst area', 'worst smoothness',
       'worst compactness', 'worst concavity', 'worst concave points',
       'worst symmetry', 'worst fractal dimension'], dtype='<U23')
```

<!-- grade_id: cell-13eab9c14ca47094, points: 0, grade: false, solution: true -->
```python
# 1 point

#load the data from the dataset into a pd df


#output a dataframe with the data of the dataset and the column names as feature names
#IT SHOULD NOT INCLUDE THE TARGET 

df_feat =pd.DataFrame(data=cancer['data'],columns=cancer['feature_names'])
# YOUR CODE HERE
df_feat.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
RangeIndex: 569 entries, 0 to 568
Data columns (total 30 columns):
 #   Column                   Non-Null Count  Dtype  
---  ------                   --------------  -----  
 0   mean radius              569 non-null    float64
 1   mean texture             569 non-null    float64
 2   mean perimeter           569 non-null    float64
 3   mean area                569 non-null    float64
 4   mean smoothness          569 non-null    float64
 5   mean compactness         569 non-null    float64
 6   mean concavity           569 non-null    float64
 7   mean concave points      569 non-null    float64
 8   mean symmetry            569 non-null    float64
 9   mean fractal dimension   569 non-null    float64
 10  radius error             569 non-null    float64
 11  texture error            569 non-null    float64
 12  perimeter error          569 non-null    float64
 13  area error               569 non-null    float64
 14  smoothness error         5
```

<!-- grade_id: cell-1c2998ddc5440385, points: 1, grade: true, solution: false -->
```python
b = True
for i in cancer['feature_names']:
    if df_feat[i].shape != (569,):
        b = False
        break
assert b
```

<!-- grade_id: cell-58bf6e127f5d9c62, points: 0, grade: false, solution: true -->
```python
# 1 point

# create a dataframe with all the targets and the column name your be 'Cancer'
# output a single column dataframe
df_target =pd.DataFrame(data=cancer['target'],columns=['Cancer'])

```

<!-- grade_id: cell-c504ee25e1401818, points: 1, grade: true, solution: false -->
```python
assert df_target['Cancer'].shape == (569,)
```

<!-- grade_id: cell-cea569421704556d, points: 0, grade: false, solution: true -->
```python
#1 point
# get the train test split, tested with random state 0 and test_size 0.2
# import if needed
from sklearn.model_selection import train_test_split
x=df_feat.to_numpy()
y=df_target['Cancer'].to_numpy()
x_train,x_test,y_train,y_test = train_test_split(x,y,test_size = 0.2,random_state =0)
# all the subsets should be numpy arrays
# TIP carefull with the shape of the Y
# YOUR CODE HERE

```

<!-- grade_id: cell-8557f6d20b5ebd4a, points: 1, grade: true, solution: false -->
```python
assert x_train.shape == (455,30)
assert x_test.shape == (114,30)
assert y_train.shape == (455,)
assert y_test.shape == (114,)
```

<!-- grade_id: cell-d569929ec688ee8f, points: 0, grade: false, solution: true -->
```python
# 2 points
#Train the support vector classifier
# import if needed
from sklearn.svm import SVC
clf=SVC()
clf.fit(x_train,y_train)

```

Output:
```
SVC()
```

<!-- grade_id: cell-741670633ab66a5b, points: 2, grade: true, solution: false -->
```python
assert clf.support_vectors_.shape == (123,30)
```

<!-- grade_id: cell-ee4e7bf6bb7b13e9, points: 0, grade: false, solution: true -->
```python
# 3 points
# predict the testing data and get the accuracy
from sklearn import metrics

# import if needed
# output preds is the array with the predictions and acc is a float between 0 and 1
preds =clf.predict(x_test)
acc=metrics.accuracy_score(y_test,preds)



```

<!-- grade_id: cell-a8b6a273210f6822, points: 2, grade: true, solution: false -->
```python
assert acc > 0.85
assert preds.shape == (114,)
```

***Everything is being classified into a single class! Use Grid Search to search for the best parameters! (Scikitlearn GridSearchCV)***


<!-- grade_id: cell-43d48d4bf75c556e, points: 0, grade: false, solution: true -->
```python
# 2 points
from sklearn.model_selection import GridSearchCV
from sklearn import svm
#Set parameters

#degrees of polynomial kernel of svc

# the parameters should be in a dict format
params = {'kernel':('linear', 'rbf'), 'C':[1, 10]}

clf = GridSearchCV(clf, params)
clf.fit(x,y)


```

Output:
```
GridSearchCV(estimator=SVC(),
             param_grid={'C': [1, 10], 'kernel': ('linear', 'rbf')})
```

```python
len(params['kernel'])
```

Output:
```
2
```

<!-- grade_id: cell-28a594d468d2e212, points: 2, grade: true, solution: false -->
```python
assert len(params) == 2
assert len(params['kernel']) == 2
```

<!-- grade_id: cell-8ac1a0df91772743, points: 0, grade: false, solution: true -->
```python
#Fit grid (Note to self: first it runs the same loop with cross-validation to find the best parameter combination. Once it has the best combination, it runs fit again on all data passed to fit (without cross-validation), to built a single new model using the best parameter setting.)
# 4 points
# add an svm parameters = {'kernel':('linear', 'rbf'), 'C':[1, 10]}
params = {'kernel':('linear', 'rbf'), 'C':[1, 10]}
svc = svm.SVC()
grid_clf = GridSearchCV(svc, params)
clf.fit(x,y)
# the SVM should be an untrained one
# import if needed 

# YOUR CODE HERE


```

Output:
```
GridSearchCV(estimator=SVC(),
             param_grid={'C': [1, 10], 'kernel': ('linear', 'rbf')})
```

<!-- grade_id: cell-085b7fafdce47092, points: 4, grade: true, solution: false -->
```python
dic = grid_clf.get_params(deep=False)
assert type(dic['estimator']) == type(svm.SVC()) 
assert dic['pre_dispatch'] == '2*n_jobs'
```

<!-- grade_id: cell-7901a09d1a9092ac, points: 0, grade: false, solution: true -->
```python
# 2 point
#Get in a list the best parameters of the grid

#output python list with the best parameters sorted


best_params = sorted(clf.cv_results_.keys())

# YOUR CODE HERE

best_params
```

Output:
```
['mean_fit_time',
 'mean_score_time',
 'mean_test_score',
 'param_C',
 'param_kernel',
 'params',
 'rank_test_score',
 'split0_test_score',
 'split1_test_score',
 'split2_test_score',
 'split3_test_score',
 'split4_test_score',
 'std_fit_time',
 'std_score_time',
 'std_test_score']
```

<!-- grade_id: cell-99bb42084e0ae89b, points: 2, grade: true, solution: false -->
```python
assert best_params[0] == 'mean_fit_time'
assert best_params[1] == 'mean_score_time'
assert best_params[2] == 'mean_test_score'
assert best_params[3] == 'param_C'
```

<!-- grade_id: cell-4eb30db85584892f, points: 0, grade: false, solution: true -->
```python
# 5 points
# predict the testing set and get the accuracy
# accuracy as always will be between 0 and 1
preds =clf.predict(x_test)
acc=metrics.accuracy_score(y_test,preds)
#acc = 
# YOUR CODE HERE

```

<!-- grade_id: cell-67663369ff5f55b0, points: 5, grade: true, solution: false -->
```python

assert acc > 0.92
assert preds[preds.argmax()] == 1
assert preds.shape == (114,)
```

<!-- grade_id: cell-b467b5e84dfff2d3, points: 0, grade: false, solution: true -->
```python
# 3 points
# get the confusion matrix
#import if needed
from sklearn.metrics import confusion_matrix
conf_mat = confusion_matrix(y_test,preds)

# YOUR CODE HERE
conf_mat


```

Output:
```
array([[45,  2],
       [ 2, 65]], dtype=int64)
```

<!-- grade_id: cell-fd7017a6eb94325a, points: 3, grade: true, solution: false -->
```python
assert conf_mat.shape == (2,2)
assert conf_mat[0][1] < 7
assert conf_mat[1][0] < 3

```

# Discussion of results

```python

```