***Support Vector Machines***
    Using SVM with Python to predict whether a breast cancer tumor is malignant or benign. Exercise uses numpy, pandas, and scikitlearn and utilizes train test split, SVM, SVC,  and GridSearch to identify the best parameters for prediction.

```python
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import seaborn as sns
%matplotlib inline

from sklearn.datasets import load_breast_cancer
from sklearn.model_selection import train_test_split
from sklearn.svm import SVC
from sklearn.metrics import accuracy_score
from sklearn.model_selection import GridSearchCV
from sklearn.svm import SVC
from sklearn.pipeline import make_pipeline
from sklearn.preprocessing import StandardScaler
```

```python

```

<!-- grade_id: cell-bbc3dc112799b91f, points: 0, grade: false, solution: true -->
```python
# 1 point

# import and load the breast cancer dataset 

# output a variable where the data is loaded straight from sklearn

cancer = 0

# YOUR CODE HERE
cancer = load_breast_cancer()
X = cancer.data
Y = cancer.target



```

<!-- grade_id: cell-89604aed63202fe9, points: 1, grade: true, solution: false -->
```python
assert cancer.data.shape == (569,30)
```

```python
print(cancer['DESCR'])
```

Output:
```
.. _breast_cancer_dataset:

Breast cancer wisconsin (diagnostic) dataset
--------------------------------------------

**Data Set Characteristics:**

    :Number of Instances: 569

    :Number of Attributes: 30 numeric, predictive attributes and the class

    :Attribute Information:
        - radius (mean of distances from center to points on the perimeter)
        - texture (standard deviation of gray-scale values)
        - perimeter
        - area
        - smoothness (local variation in radius lengths)
        - compactness (perimeter^2 / area - 1.0)
        - concavity (severity of concave portions of the contour)
        - concave points (number of concave portions of the contour)
        - symmetry
        - fractal dimension ("coastline approximation" - 1)

        The mean, standard error, and "worst" or largest (mean of the three
        worst/largest values) of these features were computed for each image,
        resulting in 30 features.  For instance, field 0 is Mean Radi
```

```python
cancer['feature_names']
```

Output:
```
array(['mean radius', 'mean texture', 'mean perimeter', 'mean area',
       'mean smoothness', 'mean compactness', 'mean concavity',
       'mean concave points', 'mean symmetry', 'mean fractal dimension',
       'radius error', 'texture error', 'perimeter error', 'area error',
       'smoothness error', 'compactness error', 'concavity error',
       'concave points error', 'symmetry error',
       'fractal dimension error', 'worst radius', 'worst texture',
       'worst perimeter', 'worst area', 'worst smoothness',
       'worst compactness', 'worst concavity', 'worst concave points',
       'worst symmetry', 'worst fractal dimension'], dtype='<U23')
```

<!-- grade_id: cell-13eab9c14ca47094, points: 0, grade: false, solution: true -->
```python
# 1 point

#load the data from the dataset into a pd df

#output a dataframe with the data of the dataset and the column names as feature names
#IT SHOULD NOT INCLUDE THE TARGET 

df_feat = 0

df_feat = pd.DataFrame(cancer.data, columns = cancer.feature_names)
df_feat.head()

```

Output:
```
mean radius  mean texture  mean perimeter  mean area  mean smoothness  \
0        17.99         10.38          122.80     1001.0          0.11840   
1        20.57         17.77          132.90     1326.0          0.08474   
2        19.69         21.25          130.00     1203.0          0.10960   
3        11.42         20.38           77.58      386.1          0.14250   
4        20.29         14.34          135.10     1297.0          0.10030   

   mean compactness  mean concavity  mean concave points  mean symmetry  \
0           0.27760          0.3001              0.14710         0.2419   
1           0.07864          0.0869              0.07017         0.1812   
2           0.15990          0.1974              0.12790         0.2069   
3           0.28390          0.2414              0.10520         0.2597   
4           0.13280          0.1980              0.10430         0.1809   

   mean fractal dimension  ...  worst radius  worst texture  worst perimeter  \
0              
```

<!-- grade_id: cell-1c2998ddc5440385, points: 1, grade: true, solution: false -->
```python
b = True
for i in cancer['feature_names']:
    if df_feat[i].shape != (569,):
        b = False
        break
assert b
```

<!-- grade_id: cell-c504ee25e1401818, points: 1, grade: true, solution: false -->
```python
assert df_target['Cancer'].shape == (569,)
```

<!-- grade_id: cell-58bf6e127f5d9c62, points: 0, grade: false, solution: true -->
```python
# 1 point

# create a dataframe with all the targets and the column name your be 'Cancer'
# output a single column dataframe
df_target = 0
# YOUR CODE HERE
df_target = pd.DataFrame()
df_target['Cancer'] = cancer.target
df_target
```

Output:
```
Cancer
0         0
1         0
2         0
3         0
4         0
..      ...
564       0
565       0
566       0
567       0
568       1

[569 rows x 1 columns]
```

<!-- grade_id: cell-cea569421704556d, points: 0, grade: false, solution: true -->
```python
#1 point
# get the train test split, tested with random state 0 and test_size 0.2
# import if needed
# all the subsets should be numpy arrays
# TIP carefull with the shape of the Y
# YOUR CODE HERE
from sklearn.model_selection import train_test_split, cross_validate

x_train, x_test, y_train, y_test = train_test_split( df_feat.values, df_target.values, test_size = 0.2, random_state = 0)

y_train, y_test = y_train.flatten(), y_test.flatten()
```

<!-- grade_id: cell-8557f6d20b5ebd4a, points: 1, grade: true, solution: false -->
```python
assert x_train.shape == (455,30)
assert x_test.shape == (114,30)
assert y_train.shape == (455,)
assert y_test.shape == (114,)
```

<!-- grade_id: cell-d569929ec688ee8f, points: 0, grade: false, solution: true -->
```python
# 2 points
#Train the support vector classifier
# import if needed

clf = 0

from sklearn.svm import SVC

clf = SVC()
clf.fit(x_train, y_train)



```

Output:
```
SVC()
```

<!-- grade_id: cell-741670633ab66a5b, points: 2, grade: true, solution: false -->
```python
assert clf.support_vectors_.shape == (123,30)
```

<!-- grade_id: cell-ee4e7bf6bb7b13e9, points: 0, grade: false, solution: true -->
```python
# 3 points
# predict the testing data and get the accuracy
# import if needed
# output preds is the array with the predictions and acc is a float between 0 and 1
preds = 0
acc = 0
# YOUR CODE HERE
from sklearn.metrics import accuracy_score

preds = clf.predict(x_test)
acc = accuracy_score(y_test, preds)
print(acc)
```

Output:
```
0.9298245614035088
```

<!-- grade_id: cell-a8b6a273210f6822, points: 2, grade: true, solution: false -->
```python
assert acc > 0.85
assert preds.shape == (114,)
```

***Everything is being classified into a single class! Use Grid Search to search for the best parameters! (Scikitlearn GridSearchCV)***


<!-- grade_id: cell-43d48d4bf75c556e, points: 0, grade: false, solution: true -->
```python
# 2 points
#Set parameters
# the parameters should be in a dict format
params = {}
params = {'kernel':('linear', 'rbf'), 'C':[1, 10]}
# YOUR CODE HERE
from sklearn.model_selection import GridSearchCV
svm = SVC()
clf = GridSearchCV(svm, params)
clf.fit(x_train, y_train)

preds = clf.predict(x_test)
acc = accuracy_score(y_test, preds)
print(acc)
```

Output:
```
0.956140350877193
```

<!-- grade_id: cell-28a594d468d2e212, points: 2, grade: true, solution: false -->
```python
assert len(params) == 2
assert len(params['kernel']) == 2
```

<!-- grade_id: cell-8ac1a0df91772743, points: 0, grade: false, solution: true -->
```python
#Fit grid (Note to self: first it runs the same loop with cross-validation to find the best parameter 
#combination. Once it has the best combination, it runs fit again on all data passed to fit 
#(without cross-validation), to built a single new model using the best parameter setting.)

# 4 points
# add an svm to the grid and train it
# the SVM should be an untrained one
# import if needed 
from sklearn import svm

svc = 0
grid_clf = 0

# YOUR CODE HERE
params = {}
params = {'kernel':('linear', 'rbf'), 'C':[1, 10]}
# YOUR CODE HERE
from sklearn.model_selection import GridSearchCV

svc = svm.SVC()
grid_clf = GridSearchCV(svc, params)
grid_clf.fit(x_train, y_train)

preds = grid_clf.predict(x_test)
acc = accuracy_score(y_test, preds)
print(acc)

# grid_clf = GridSearchCV(estimator=svc, 
#                    param_grid=params)
```

Output:
```
0.956140350877193
```

<!-- grade_id: cell-085b7fafdce47092, points: 4, grade: true, solution: false -->
```python
dic = grid_clf.get_params(deep=False)
assert type(dic['estimator']) == type(svm.SVC()) 
assert dic['pre_dispatch'] == '2*n_jobs'
```

<!-- grade_id: cell-7901a09d1a9092ac, points: 0, grade: false, solution: true -->
```python
# 2 point
#Get in a list the best parameters of the grid
#output python list with the best parameters sorted

best_params = sorted(clf.cv_results_.keys())


```

<!-- grade_id: cell-99bb42084e0ae89b, points: 2, grade: true, solution: false -->
```python
assert best_params[0] == 'mean_fit_time'
assert best_params[1] == 'mean_score_time'
assert best_params[2] == 'mean_test_score'
assert best_params[3] == 'param_C'
```

<!-- grade_id: cell-4eb30db85584892f, points: 0, grade: false, solution: true -->
```python
# 5 points
# predict the testing set and get the accuracy
# accuracy as always will be between 0 and 1
preds = 0
acc = 0
# YOUR CODE HERE
preds = grid_clf.predict(x_test)
acc = accuracy_score(y_test, preds)
print(acc)

```

Output:
```
0.956140350877193
```

<!-- grade_id: cell-67663369ff5f55b0, points: 5, grade: true, solution: false -->
```python

assert acc > 0.92
assert preds[preds.argmax()] == 1
assert preds.shape == (114,)
```

<!-- grade_id: cell-b467b5e84dfff2d3, points: 0, grade: false, solution: true -->
```python
# 3 points
# get the confusion matrix
#import if needed
conf_mat = []

# YOUR CODE HERE

from sklearn.metrics import plot_confusion_matrix
from sklearn.metrics import confusion_matrix

conf_mat = confusion_matrix(preds, y_test)
#conf_mat = np.transpose(conf_mat)
```

<!-- grade_id: cell-fd7017a6eb94325a, points: 3, grade: true, solution: false -->
```python
assert conf_mat.shape == (2,2)
assert conf_mat[0][1] < 7
assert conf_mat[1][0] < 3

```

# Discussion of results

```python

```